# ENG 209 Midterm - Problèmes ouverts

```python
from typing import Callable, Any
from dataclasses import dataclass, field
from datetime import datetime
from pprint import pprint
from collections import defaultdict
try:
    import numpy as np
except:
    print("Run 'pip install numpy' in a terminal")
```

<div style="font-size: 120%">

___Les problèmes ci-dessous sont à résoudre individuellement.  
Vous n’avez pas accès à Copilot, ChatGPT ou à tout autre outil de génération de code.  
Vous pouvez vous référer à votre résumé/formulaire.  
La consultation des séries précédentes est interdite.___

* N’écrivez du code **que** dans les cellules avec `### YOUR CODE HERE`
* **Supprimez** les lignes `raise NotImplementedError()`
* **N’ajoutez pas** de cellules et ne détruisez pas les cellules vides.
</div>

<hr>

<div style="font-size: 120%">

___The problems below must be solved individually.  
You do not have access to Copilot, ChatGPT, or any other code generation tool.  
You may refer to your summary/form sheet.  
Consulting previous exercises is not allowed.___

* Write code **only** in the cells marked with `### YOUR CODE HERE`
* **Delete** the lines `raise NotImplementedError()`
* **Do not add** or remove any cells, including empty ones.
</div>

---

## Problème 1

<table width="100%;">
<tr>
<td width="50%" style="vertical-align: top; border: none; border-right: 1px solid lightgrey; padding-right: 10px; align: center;">

Implémentez la fonction `reduce` ci-dessous. Elle accepte une liste de `int`s et une fonction `f` qui prend deux `int`s et retourne un `int`, et un élément initial `initial`. La valeur qu'elle retourne est le résultat de l'application répétée de `f` aux éléments de la liste, en partant de la gauche, et en considérant `initial` comme valeur de départ.

Par exemple, si `values = [1, 2]`, `f = lambda x, y: x + y` (simplement, l'addition de deux `int`s) et `initial = 0`, alors `reduce(values, f, initial)` doit retourner `3`, soit la somme des éléments de la liste, car le calcul est effectué comme suit: `f(f(initial, 1), 2) = f(f(0, 1), 2) = ((0 + 1) + 2) = 3`.

</td>
<td width="50%" style="vertical-align: top; border: none; align: center">

Implement the `reduce` function below. It takes a list of `int`s, a function `f` that takes two `int`s and returns an `int`, and an initial element `initial`. The value it returns is the result of repeatedly applying `f` to the elements of the list from left to right, starting with `initial` as the initial value.

For example, if `values = [1, 2]`, `f = lambda x, y: x + y` (simply the addition of two `int`s), and `initial = 0`, then `reduce(values, f, initial)` should return `3`, i.e. the sum of the list elements, since the computation proceeds as follows: `f(f(initial, 1), 2) = f(f(0, 1), 2) = ((0 + 1) + 2) = 3`.

</td>
</tr>
</table>

<!-- grade_id: reduce, points: 0, grade: false, solution: true -->
```python
def reduce(values: list[int], f: Callable[[int, int], int], initial: int) -> int:
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
# devrait afficher/should display 10
reduce([1, 2, 3, 4], lambda x, y: x + y, 0)
```

```python
# devrait afficher/should display 24
reduce([1, 2, 3, 4], lambda x, y: x * y, 1)
```

<!-- grade_id: reduce_tests_1, points: 1, grade: true, solution: false -->
```python
# Tests - ne pas modifier/do not modify
assert reduce([1, 2, 3, 4], lambda x, y: x + y, 0) == 10, "addition does not work"
assert reduce([1, 2, 3, 4], lambda x, y: x * y, 1) == 24, "multiplication does not work"
```

<!-- grade_id: reduce_tests_2, points: 1, grade: true, solution: false -->
```python
# Tests - ne pas modifier/do not modify
```

---

## Problème 2

<table width="100%;">
<tr>
<td width="50%" style="vertical-align: top; border: none; border-right: 1px solid lightgrey; padding-right: 10px; align: center;">

Complétez la fonction `to_dict` ci-dessous, **sans modifier** la classe `Record`. N'**utilisez pas `defaultdict`**.

La fonction accepte une liste de tuples `(str, int)` représentant des heures travaillées par des employés identifiés par une chaîne de caractères (leur nom).

La fonction doit retourner un dictionnaire où chaque clé est le nom d'un employé et la valeur associée est une instance de la classe `Record`, qui indique le nom de l'employé, la liste de ses heures travaillées et le total des heures.

Par exemple, si `pairs = [('a', 1), ('b', 2), ('a', 3), ('b', 2)]`, le dictionnaire retourné devra être:

````
{'a': Record(name='a', values=[1, 3], total=4),
 'b': Record(name='b', values=[2, 2], total=4)}
````

</td>
<td width="50%" style="vertical-align: top; border: none; align: center">

Complete the `to_dict` function below **without modifying** the `Record` class. Do **not use `defaultdict`**.

The function takes a list of tuples `(str, int)` representing hours worked by employees identified by a string (their name).

It must return a dictionary where each key is an employee’s name and the associated value is an instance of the `Record` class, containing the employee’s name, the list of their worked hours, and the total number of hours.

For example, if `pairs = [('a', 1), ('b', 2), ('a', 3), ('b', 2)]`,  
the returned dictionary should be:

````
{'a': Record(name='a', values=[1, 3], total=4),
 'b': Record(name='b', values=[2, 2], total=4)}
````

</td>
</tr>
</table>



<!-- grade_id: to_dict, points: 0, grade: false, solution: true -->
```python
@dataclass
class Record:
    name: str
    values: list[int]
    total: int


def to_dict(pairs: list[tuple[str, int]]) -> dict[str, Record]:
    # YOUR CODE HERE
    raise NotImplementedError()

```

```python
# devrait afficher/should display:
# {'a': Record(name='a', values=[1, 3], total=4),
#  'b': Record(name='b', values=[2, 2], total=4)}
to_dict([('a', 1), ('b', 2), ('a', 3), ('b', 2)])
```

<!-- grade_id: to_dict_tests_1, points: 1, grade: true, solution: false -->
```python
# Tests - ne pas modifier/do not modify
assert to_dict([("a", 1), ("b", 2), ("a", 3)]) == {'a': Record(name='a', values=[1, 3], total=4), 'b': Record(name='b', values=[2], total=2)}, "to_dict does not work"
```

<!-- grade_id: to_dict_tests_2, points: 1, grade: true, solution: false -->
```python
# Tests - ne pas modifier/do not modify
```

---
## Problème 3

<table width="100%">
<tr>
<td width="50%" style="vertical-align: top; border: none; border-right: 1px solid lightgrey; padding-right: 10px; align: center;">

Complétez la classe `Contact` ci-dessous en suivant les instructions suivantes.

 1. Donnez-lui un champ en plus appelé `friends`, qui est une liste de `Contact`.  
    Pour le type, utilisez `list["Contact"]` (notez les guillemets) plutôt que `list[Contact]` pour éviter les problèmes de dépendances circulaires.  
    Ce champ doit être initialisé par défaut à une nouvelle liste vide pour chaque nouvelle instance créée.
 2. Ajoutez une propriété `full_name` qui retourne le nom complet du contact.  
    On considère que le `first_name` et le `last_name` sont toujours définis, mais que `middle_name` peut être vide (`""`).  
    Le `full_name` doit être la concaténation des trois noms, séparés par un espace, mais sans espaces superflus. Par exemple, si `first_name="John"`, `middle_name=""` et `last_name="Doe"`, alors `full_name` doit retourner `"John Doe"`.

    Faites en sorte que le `full_name` soit calculé à la demande et non stocké dans un champ et s'obtienne avec la syntaxe `p.full_name` (sans parenthèses), et non `p.full_name()`, si `p` est une instance de `Contact`.
 3. Ajoutez une méthode `filtered_friends` qui accepte un prédicat (une fonction qui prend un `Contact` et retourne un booléen, de type `Callable[['Contact'], bool]`) et retourne une liste des amis qui satisfont ce prédicat.  
    Implémentez cette méthode **en une seule ligne**.

</td>
<td width="50%" style="vertical-align: top; border: none; align: center">

Complete the `Contact` class below according to the following instructions.

1. Add an extra field called `friends`, which is a list of `Contact`.  
   For the type annotation, use `list["Contact"]` (note the quotes) instead of `list[Contact]` to avoid circular dependency issues.  
   This field must be initialized by default to a new empty list for each new instance created.

2. Add a property `full_name` that returns the contact’s full name.  
   The `first_name` and `last_name` are always defined, but `middle_name` may be empty (`""`).  
   The `full_name` should be the concatenation of the three names, separated by spaces, but without extra spaces.  
   For example, if `first_name="John"`, `middle_name=""`, and `last_name="Doe"`, then `full_name` should return `"John Doe"`.

   The `full_name` must be computed on demand (not stored in a field) and accessed using the syntax `p.full_name` (without parentheses), not `p.full_name()`, if `p` is an instance of `Contact`.

3. Add a method `filtered_friends` that takes a predicate (a function that accepts a `Contact` and returns a boolean, of type `Callable[['Contact'], bool]`) and returns a list of friends that satisfy this predicate.  
   Implement this method **in a single line**.
<td>
</tr>
</table>

<!-- grade_id: contact, points: 0, grade: false, solution: true -->
```python
@dataclass
class Contact:
    first_name: str
    middle_name: str
    last_name: str
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
p1 = Contact("John", "B.", "Robertson")
p2 = Contact("Joyce", "Jenna", "Brown")
p3 = Contact("James", "", "Kerringan")

p1.friends.append(p2)
p1.friends.append(p3)

print(p1.full_name)  # John B. Robertson
print(p2.full_name)  # Joyce Jenna Brown
print(p3.full_name)  # James Kerringan
```

<!-- grade_id: contact_tests_1, points: 1, grade: true, solution: false -->
```python
# Tests - ne pas modifier/do not modify
assert Contact("John", "B.", "Robertson").full_name == "John B. Robertson", "full_name does not work"
```

<!-- grade_id: contact_tests_2, points: 1, grade: true, solution: false -->
```python
# Tests - ne pas modifier/do not modify
```

<!-- grade_id: contact_tests_3, points: 1, grade: true, solution: false -->
```python
# Tests - ne pas modifier/do not modify
```

<!-- grade_id: contact_tests_4, points: 1, grade: true, solution: false -->
```python
# Tests - ne pas modifier/do not modify
```

---
## Problème 4

<table width="100%;">
<tr>
<td width="50%" style="vertical-align: top; border: none; border-right: 1px solid lightgrey; padding-right: 10px;">

L’angle $\theta$ entre deux vecteurs $\vec{u}$ et $\vec{v}$ dans un espace euclidien est défini par la formule suivante:

</td>
<td width="50%" style="vertical-align: top; border: none;">

The angle $\theta$ between two vectors $\vec{u}$ and $\vec{v}$ in a Euclidean space is defined by the following formula:
</td>
</tr>
</table>

$$\theta = \arccos\left(\frac{\vec{u} \cdot \vec{v}}{\|\vec{u}\|\|\vec{v}\|}\right)$$


<table width="100%">
<tr>
<td width="50%" style="vertical-align: top; border: none; border-right: 1px solid lightgrey; padding-right: 10px;">

Voici une implémentation de cette formule en Python, utilisant NumPy:

</td>
<td width="50%" style="vertical-align: top; border: none;">

Here is an implementation of this formula in Python using NumPy:
</td>
</tr>
</table>



```python
def angle_between(v1: np.ndarray, v2: np.ndarray) -> float:
    numer = np.dot(v1, v2)
    denom = np.linalg.norm(v1) * np.linalg.norm(v2)
    if denom == 0:
        raise ValueError("One of the vectors has zero magnitude.")
    angle_radians = np.arccos(numer / denom)
    return float(np.degrees(angle_radians))
```

<table width="100%;">
<tr>
<td width="50%" style="vertical-align: top; border: none; border-right: 1px solid lightgrey; padding-right: 10px;">

Complétez la fonction `random_vector(num_dimensions: int) -> np.ndarray`, qui génère un vecteur aléatoire de `num_dimensions` dimensions, où chaque composante est un nombre flottant dans l'intervalle $[-1, 1)$.

Pour cela, utilisez d'abord `np.random.rand(num_dimensions)`. Cela génère un vecteur dont chaque composante est un nombre aléatoire dans l'intervalle $[0, 1)$. Modifiez-le ensuite avec des opérations arithmétiques pour que les composantes soient dans l'intervalle $[-1, 1)$ afin de couvrir toutes les directions possibles.

</td>
<td width="50%" style="vertical-align: top; border: none;">

Complete the function `random_vector(num_dimensions: int) -> np.ndarray`, which generates a random vector of `num_dimensions` dimensions, where each component is a floating-point number in the interval $[-1, 1)$.

To do this, first use `np.random.rand(num_dimensions)`.
This generates a vector where each component is a random number in the interval $[0, 1)$.
Then modify it using arithmetic operations so that the components fall within the interval $[-1, 1)$, in order to cover all possible directions.

</td>
</tr>
</table>

<!-- grade_id: random_vector, points: 0, grade: false, solution: true -->
```python
def random_vector(num_dimensions: int) -> np.ndarray:
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
print(random_vector(1))
print(random_vector(2))
print(random_vector(3))
print(random_vector(10))
```

<!-- grade_id: random_vector_tests_1, points: 1, grade: true, solution: false -->
```python
# Tests - ne pas modifier/do not modify
assert random_vector(3).shape == (3,)
```

<table width="100%;">
<tr>
<td width="50%" style="vertical-align: top; border: none; border-right: 1px solid lightgrey; padding-right: 10px;">

Complétez la fonction `sample_mean_angle_between_random_vectors` ci-dessous afin qu'elle calcule une estimation de l'angle moyen entre deux vecteurs aléatoires.

Pour cela, elle doit générer `num_samples` paires de vecteurs aléatoires à `num_dimensions` composantes (en appelant `random_vector`), calculer l'angle entre chaque paire en utilisant la fonction `angle_between`, et retourner la moyenne de ces angles.

</td>
<td width="50%" style="vertical-align: top; border: none;">

Complete the `sample_mean_angle_between_random_vectors` function below so that it computes an estimate of the mean angle between two random vectors.  

To do this, it should generate `num_samples` pairs of random vectors with `num_dimensions` components (by calling `random_vector`), compute the angle between each pair using the `angle_between` function, and return the mean of these angles.

</td>
</tr>
</table>

<!-- grade_id: sample_mean, points: 0, grade: false, solution: true -->
```python
def sample_mean_angle_between_random_vectors(num_samples: int, num_dimensions: int) -> float:
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: sample_mean_tests_1, points: 2, grade: true, solution: false -->
```python
# Tests - ne pas modifier/do not modify
sample_mean_angle_between_random_vectors(1000, 3)
```